# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`. Additionally, there are several countries with Numeric digits in their name. Make sure to remove these as well, e.g. `'Italy9'` should be `'Italy'`. 

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    Energy = Energy[0:15]
    print(Energy.index)
    return Energy
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

Output:
```
Index(['China', 'United States', 'Japan', 'United Kingdom',
       'Russian Federation', 'Canada', 'Germany', 'India', 'France',
       'South Korea', 'Italy', 'Spain', 'Iran', 'Australia', 'Brazil'],
      dtype='object')
Index(['China', 'United States', 'Japan', 'United Kingdom',
       'Russian Federation', 'Canada', 'Germany', 'India', 'France',
       'South Korea', 'Italy', 'Spain', 'Iran', 'Australia', 'Brazil'],
      dtype='object')
```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    #alll = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='outer').join(ScimEn.set_index('Country'),how ='outer')

    allentries = len(Energy)+len(GDP)+len(ScimEn)

    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    af = len(Energy)
    print(allentries-af)
    return allentries-af
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

Output:
```
522
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    #GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    #GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    Energy = Energy[0:15]
    GDP =    Energy[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    #print(GDP.dtypes)
    
    #GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'] = pd.to_numeric(GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'],errors='coerce')
    GDP['avgGDP'] = GDP.apply(np.nanmean, axis=1)
    GDP.sort_values(by=['avgGDP'],ascending = False, inplace = True)
    #GDP = GDP[0:15]
    #GDP.set_index('Country Name',inplace = True)

    avgGDP = GDP['avgGDP']
    #print(avgGDP[0:25])

    #avgGDP.set_index('Country Name',inplace = True)
    
    if isinstance(avgGDP, pd.DataFrame):
        print("my_series is a pandas series")

    print(avgGDP)
    return avgGDP
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

Output:
```
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
Iran                  4.441558e+11
Name: avgGDP, dtype: float64
```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    #GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    Energy = Energy[0:15]
    GDP =    Energy[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    #print(GDP.dtypes)
    #GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'] = pd.to_numeric(GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'],errors='coerce')
    GDP['avgGDP'] = GDP.apply(lambda x: x.mean(), axis=1)
    GDP.sort_values(by=['avgGDP'],ascending = False, inplace = True)
    #avgGDP = GDP['avgGDP']

    GDP = GDP.index[5]

    #GDP.set_index('Country Name',inplace = True)
    #print(Energy)
    #print(Energy.loc[GDP]['2015']-Energy.loc[GDP]['2006'])
    return Energy.loc[GDP]['2015']-Energy.loc[GDP]['2006']
answer_four()
```

Output:
```
246702696075.3999
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    Energy = Energy[0:15]['Energy Supply per Capita']
    print(Energy)
    mean = Energy.mean()
    print(mean)
    return mean
answer_five()
```

Output:
```
China                  93.0
United States         286.0
Japan                 149.0
United Kingdom        124.0
Russian Federation    214.0
Canada                296.0
Germany               165.0
India                  26.0
France                166.0
South Korea           221.0
Italy                 109.0
Spain                 106.0
Iran                  119.0
Australia             231.0
Brazil                 59.0
Name: Energy Supply per Capita, dtype: float64
157.6
157.6
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)

    Energy = pd.to_numeric(Energy[0:15]['% Renewable'])
    maximum = (Energy.idxmax(),Energy.max()) 
    #GDP =    Energy[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    #print(GDP.dtypes)
    #GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'] = pd.to_numeric(GDP['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'],errors='coerce')
    #GDP['avgGDP'] = GDP.apply(lambda x: x.mean(), axis=1)
    #GDP.sort_values(by=['avgGDP'],ascending = False, inplace = True)   
    print(maximum)
    
    return maximum
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

Output:
```
('Brazil', 69.64803)
('Brazil', 69.64803)
```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)

    Energy = pd.to_numeric(Energy[0:15]['Self-citations']/Energy[0:15]['Citations'])
    maximum2 = (Energy.idxmax(),Energy.max())
    print(maximum2)
    
    
    
    return maximum2
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

Output:
```
('China', 0.6893126179389422)
('China', 0.6893126179389422)
```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)

    Energy = pd.to_numeric(Energy[0:15]['Energy Supply']/Energy[0:15]['Energy Supply per Capita'])
    Energy.sort_values( inplace = True)

    print(Energy,Energy.index[12])
    
    
    
    return Energy.index[12]
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

Output:
```
Australia             2.331602e+07
Canada                3.523986e+07
Spain                 4.644340e+07
South Korea           4.980543e+07
Italy                 5.990826e+07
France                6.383735e+07
United Kingdom        6.387097e+07
Iran                  7.707563e+07
Germany               8.036970e+07
Japan                 1.274094e+08
Russian Federation    1.435000e+08
Brazil                2.059153e+08
United States         3.176154e+08
India                 1.276731e+09
China                 1.367645e+09
dtype: float64 United States
```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)
    Energy = Energy[0:15]
    Energy['population'] = pd.to_numeric(Energy['Energy Supply']/Energy['Energy Supply per Capita'])
    Energy['citation per person'] =     Energy['Citable documents']/Energy['population'] 
    Energy = Energy[['Energy Supply per Capita','citation per person']]
    print(Energy.corr().loc['Energy Supply per Capita']['citation per person'])
    #Energy.sort_values( inplace = True)
    
    
    return Energy.corr().loc['Energy Supply per Capita']['citation per person']
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

Output:
```
0.7940010435442943
0.7940010435442943
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)
    Energy = pd.to_numeric(Energy[0:15]['% Renewable'])
    HighRenew = Energy#.sort_values(ascending = False)
    median = HighRenew.median()
    HighRenew = HighRenew.where(HighRenew>= median, 0)
    HighRenew = HighRenew.where(HighRenew< median, 1)

    print(HighRenew,Energy)
    
    
    
    
    return HighRenew
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

Output:
```
China                 1.0
United States         0.0
Japan                 0.0
United Kingdom        0.0
Russian Federation    1.0
Canada                1.0
Germany               1.0
India                 0.0
France                1.0
South Korea           0.0
Italy                 1.0
Spain                 1.0
Iran                  0.0
Australia             0.0
Brazil                1.0
Name: % Renewable, dtype: float64 China                 19.754910
United States         11.570980
Japan                 10.232820
United Kingdom        10.600470
Russian Federation    17.288680
Canada                61.945430
Germany               17.901530
India                 14.969080
France                17.020280
South Korea            2.279353
Italy                 33.667230
Spain                 37.968590
Iran                   5.707721
Australia             11.810810
Brazil                69.648030
Name: % Renewable, dtype: float64
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                      'United States':'North America', 
                      'Japan':'Asia', 
                      'United Kingdom':'Europe', 
                      'Russian Federation':'Europe', 
                      'Canada':'North America', 
                      'Germany':'Europe', 
                      'India':'Asia',
                      'France':'Europe', 
                      'South Korea':'Asia', 
                      'Italy':'Europe', 
                      'Spain':'Europe', 
                      'Iran':'Asia',
                      'Australia':'Australia', 
                      'Brazil':'South America'}
    #Continent = pd.DataFrame.from_dict(ContinentDict, orient='index')
    #print(Continent)
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)
    Energy = Energy[0:15]
    Energy['population'] = pd.to_numeric(Energy['Energy Supply']/Energy['Energy Supply per Capita'])
    Energy = Energy['population']
    #Energy = Energy.join( Continent)
    grouped = Energy
    size = grouped.groupby(by = ContinentDict).count()
    summ = grouped.groupby(by = ContinentDict).sum()
    mean = grouped.groupby(by = ContinentDict).mean()
    std = grouped.groupby(by = ContinentDict).std()
    alll =pd.concat([size,summ,mean,std],axis = 1,keys = ['size','sum','mean','std'])

    print(alll)
    
    
    
    return alll
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

Output:
```
size           sum          mean           std
Asia              5  2.898666e+09  5.797333e+08  6.790979e+08
Australia         1  2.331602e+07  2.331602e+07           NaN
Europe            6  4.579297e+08  7.632161e+07  3.464767e+07
North America     2  3.528552e+08  1.764276e+08  1.996696e+08
South America     1  2.059153e+08  2.059153e+08           NaN
               size           sum          mean           std
Asia              5  2.898666e+09  5.797333e+08  6.790979e+08
Australia         1  2.331602e+07  2.331602e+07           NaN
Europe            6  4.579297e+08  7.632161e+07  3.464767e+07
North America     2  3.528552e+08  1.764276e+08  1.996696e+08
South America     1  2.059153e+08  2.059153e+08           NaN
               size           sum          mean           std
Asia              5  2.898666e+09  5.797333e+08  6.790979e+08
Australia         1  2.331602e+07  2.331602e+07           NaN
Europe            6  4.579297e+08  7.632161e+07  3.464767e+07
North America     2  3.
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERe
    ContinentDict  = {'China':'Asia', 
                      'United States':'North America', 
                      'Japan':'Asia', 
                      'United Kingdom':'Europe', 
                      'Russian Federation':'Europe', 
                      'Canada':'North America', 
                      'Germany':'Europe', 
                      'India':'Asia',
                      'France':'Europe', 
                      'South Korea':'Asia', 
                      'Italy':'Europe', 
                      'Spain':'Europe', 
                      'Iran':'Asia',
                      'Australia':'Australia', 
                      'Brazil':'South America'}
    #Continent = pd.DataFrame.from_dict(ContinentDict, orient='index')
    #print(Continent)
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)
    Energy = Energy[0:15]
    Energy['population'] = pd.to_numeric(Energy['Energy Supply']/Energy['Energy Supply per Capita'])
    Energy = Energy['% Renewable']
    #Energy = Energy.join( Continent)
    grouped = Energy
    size = grouped.groupby(by = ContinentDict).count()
    print(size)
    return size

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 25, "Q12: Wrong result numbers!"

```

Output:
```
Asia             5
Australia        1
Europe           6
North America    2
South America    1
Name: % Renewable, dtype: int64
Asia             5
Australia        1
Europe           6
North America    2
South America    1
Name: % Renewable, dtype: int64
AssertionError                            Traceback (most recent call last)
Cell In [26], line 3
      1 assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"
----> 3 assert len(answer_twelve()) == 25, "Q12: Wrong result numbers!"

AssertionError: Q12: Wrong result numbers!
```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows = 16,skipfooter = 38)
    Energy = Energy.drop(columns = ['Unnamed: 0','Unnamed: 1']).rename(columns = {'Unnamed: 2':'Country','Renewable Electricity Production':'% Renewable'} )
    Energy['Country'] = Energy.Country.str.replace('[0-9]','', regex=True) 
    Energy['Country'] = Energy.Country.str.replace('\ \(.*\)','', regex=True) 
    Energy = Energy.replace(["Republic of Korea","United States of America","United Kingdom of Great Britain and Northern Ireland","China, Hong Kong Special Administrative Region"] ,["South Korea", "United States","United Kingdom","Hong Kong"])
    Energy['Energy Supply'][1:] = pd.to_numeric(Energy['Energy Supply'][1:],errors='coerce')* 1000000
    Energy['Energy Supply'][0] = 'gigajoules'
    Energy['Energy Supply per capita'] = pd.to_numeric(Energy['Energy Supply per capita'][1:],errors='coerce')
    GDP = pd.read_csv('assets/world_bank.csv',skiprows = 4)
    GDP = GDP.replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"})
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    Energy = Energy.set_index('Country').join(GDP.set_index('Country Name'),how ='inner').join(ScimEn.set_index('Country'),how ='inner')
    Energy.sort_values(by=['Rank'], inplace = True)
    Energy = Energy[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    Energy.rename(columns = {'Energy Supply per capita':'Energy Supply per Capita'},inplace = True)
    #print(Energy.dtypes)
    Energy = Energy[0:15]
    Energy['population'] = pd.to_numeric(Energy['Energy Supply']/Energy['Energy Supply per Capita'])
    PopEst = Energy
    PopEst['pop'] = PopEst['population'].map('{:,}'.format)
    PopEst = PopEst['pop']
    PopEst.name = 'PopEst'
    print(PopEst)
    return PopEst
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

Output:
```
China                 1,367,645,161.2903225
United States          317,615,384.61538464
Japan                  127,409,395.97315437
United Kingdom         63,870,967.741935484
Russian Federation            143,500,000.0
Canada                  35,239,864.86486486
Germany                 80,369,696.96969697
India                 1,276,730,769.2307692
France                  63,837,349.39759036
South Korea            49,805,429.864253394
Italy                  59,908,256.880733944
Spain                    46,443,396.2264151
Iran                    77,075,630.25210084
Australia              23,316,017.316017315
Brazil                 205,915,254.23728815
Name: PopEst, dtype: object
China                 1,367,645,161.2903225
United States          317,615,384.61538464
Japan                  127,409,395.97315437
United Kingdom         63,870,967.741935484
Russian Federation            143,500,000.0
Canada                  35,239,864.86486486
Germany                 80,369,696.96969697
Indi
```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```