# Assignment 2 - Network Connectivity

In this assignment you will go through the process of importing and analyzing an internal email communication network between employees of a mid-sized manufacturing company. 
Each node represents an employee and each directed edge between two nodes represents an individual email. The left node represents the sender and the right node represents the recipient. We will also store the timestamp of each email. 

```python
import networkx as nx

#!head assets/email_network.txt
```

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
def answer_one():
    G1 = nx.read_edgelist('assets/email_network.txt',nodetype=int,
    data=(('time',float),),delimiter ='\t',create_using=nx.MultiDiGraph())
      #nx.draw(G1, with_labels = True)
      #plt.show()
        # YOUR CODE HERE
    return G1
```

<!-- grade_id: cell-483a56c89f9db231, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recipient` pairs of employees are there in the network such that `sender` sent at least one email to `recipient`? Note that even if a `sender` sent multiple messages to a `recipient`, they should only be counted once. You should **not** exclude cases where an employee sent emails to themselves from this [email] count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recipient` pairs).*

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    
    #print(len(set(G1.edges())))
    #print(G1.edges.data())
    G1 = ans_one
    return G1.number_of_nodes(),len(set(G1.edges()))
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    G1 = ans_one
    print(nx.is_strongly_connected(G1))
    print(nx.is_weakly_connected(G1))

    
    return nx.is_strongly_connected(G1),nx.is_weakly_connected(G1)
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

Output:
```
False
True
```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    G1 = ans_one
    length = 0
    components =nx.weakly_connected_components(G1) 
    for i in components:
        if length < len(i):
            length = len(i)

    return length
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    G1 = ans_one
    length = 0
    components =nx.strongly_connected_components(G1) 
    for i in components:
        if length < len(i):
            length = len(i)
    return length
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
        # YOUR CODE HERE
    G1 = ans_one
    components =nx.strongly_connected_components(G1) 
    nodes = [c for c in sorted(nx.strongly_connected_components(G1), key=len, reverse=True)]
    print(nodes[0])
    common_nodes = set(nodes[0])
    print(common_nodes)
    G_sc = G1.subgraph(common_nodes)
    return G_sc
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

Output:
```
{1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 103, 104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 115, 117, 118, 119, 120, 121, 122, 123, 124, 128, 129, 134, 149, 151}
{1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 103, 104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 115, 117, 118, 119, 120,
```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    G1 = ans_six
    return nx.average_shortest_path_length(G1)
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    G1 = ans_six


    return nx.diameter(G1)
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    G1 = ans_six
    aaa =nx.eccentricity(G1)
    nodes = [str(c)  for c in aaa if aaa[c]==nx.diameter(G1)]
    print(nodes)
    return set(nodes)
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
```

Output:
```
['97', '129', '134']
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    G1 = ans_six
    aaa =nx.eccentricity(G1)
    length = nx.radius(G1) 
    print(length)
    nodes = [str(c)  for c in aaa if aaa[c]==length]
    print(nodes)
    return set(nodes)
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

Output:
```
1
[38]
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    #print(nx.all_shortest_paths(set0))
    #short_paths = nx.shortest_path(G1,set0)
    #print(G1.nodes())
    G1,set0 = ans_six,ans_nine
    source = [int(c) for c in set0]
    #print(source)
    #source = 97
    result = [(c,nx.shortest_path(G1, source=c)) for c in source]
    nodes = [( source[i],j) for i in range(len(set0)) for j in G1.nodes() if len(result[i][1][j]) == 4]
    counts = {value: sum(1 for tup in nodes if tup[0] == value) for value in set(tup[0] for tup in nodes)}
    print(counts)


    return str(97),counts[97]
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

Output:
```
{97: 63, 129: 1, 134: 2}
```

### Question 12

Suppose you want to prevent communication flow from the node that you found in question 11 to node 10. What is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or 10)? 

*This function should return an integer.*

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    G_sc = ans_six
    print(nx.node_connectivity(G_sc,97,10))


    return nx.node_connectivity(G_sc,97,10)
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()

```

Output:
```
1
```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE

    G_sc = ans_six
    G_un = G_sc.to_undirected()
    print(type(G_un))
    return nx.Graph(G_un)
```

<!-- grade_id: cell-d1c0627a327cd774, points: 1, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

Output:
```
<class 'networkx.classes.multigraph.MultiGraph'>
```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clustering).*     
*Note: DO NOT round up your answer.*

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    # YOUR CODE HERE
    G_un = ans_thirteen
    print(nx.transitivity(G_un),nx.average_clustering(G_un))




    return nx.transitivity(G_un),nx.average_clustering(G_un)
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```

Output:
```
0.570111160700385 0.6975272437231418
```

```python

```